<a href="https://colab.research.google.com/github/JayAiJoe/CSMODEL-Project-1/blob/main/challenge0.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Jay-ai Locsin"
COLLABORATORS = ""
```

---

# Challenge 00: Suggested Reading


This is a self-reported challenge. You have to do all the Parts if you want to complete it. Appendices are optional.


## Part 1: Linear algebra
Assign the variable below the value `True` if you completed 3Blue1Brown's [linear algebra sequence](https://www.3blue1brown.com/topics/linear-algebra) or if you already have the prerequisite background:


<!-- grade_id: lin-alg-ans, points: 0, grade: false, solution: true -->
```python
linear_algebra_completed = True # YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: lin-alg-done, points: 1, grade: true, solution: false -->
```python
assert linear_algebra_completed
```

## Part 2: Calculus

Assign the variable below the value `True` if you completed 3Blue1Brown's [calculus sequence](https://www.3blue1brown.com/topics/calculus) or if you already have the prerequisite background:

<!-- grade_id: calc-ans, points: 0, grade: false, solution: true -->
```python
calculus_completed = True # YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: calc-done, points: 1, grade: true, solution: false -->
```python
assert calculus_completed
```

## Part 3: Python

Assign the variable below the value `True` if you completed the **first nine modules** of the [official Python tutorial](https://docs.python.org/3/tutorial/) or if you already have the prerequisite background:

<!-- grade_id: python-ans, points: 0, grade: false, solution: true -->
```python
python_completed = True # YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: python-done, points: 1, grade: true, solution: false -->
```python
assert python_completed
```

## Appendix: An excerpt from _Proofs and Refutations_, by Imre Lakatos

Assign the variable below the value `True` if you read the [excerpt](https://math.berkeley.edu/~kpmann/Lakatos.pdf) linked in the participant guide:

<!-- grade_id: appendix-ans, points: 0, grade: false, solution: true -->
```python
appendix_completed = False # YOUR CODE HERE
# raise NotImplementedError()
```

<!-- grade_id: appendix-test, points: 1, grade: true, solution: false -->
```python
assert appendix_completed
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-11-4a5a5fc4545c> in <cell line: 1>()
----> 1 assert appendix_completed

AssertionError:
```

```python

```